
<head> 
<!-- Javascript -->
<script type="text/javascript" src="static/js/runestonebase.js"></script>
<script type="text/javascript" src="static/js/jquery.js"></script>
<script type="text/javascript" src="static/js/pytutor.js"></script>
<script type="text/javascript" src="static/js/jquery-ui-1.10.3.custom.min.js"></script>
<script src="skulpt/dist/skulpt.min.js" type="text/javascript"></script> 
<script src="skulpt/dist/skulpt-stdlib.js" type="text/javascript"></script>
<script src="static/js/editor.js" type="text/javascript"></script> 
<script src="static/js/activecode.js" type="text/javascript"></script>
<script src="static/js/timed_activecode.js" type="text/javascript"></script>
<script src="codemirror/lib/codemirror.js" type="text/javascript"></script>
<script src="codemirror/addon/active-line.js" type="text/javascript"></script>
<script src="codemirror/addon/search.js" type="text/javascript"></script>
<script src="codemirror/addon/match-highlighter.js" type="text/javascript"></script>
<script src="codemirror/addon/matchesonscrollbar.js" type="text/javascript"></script>
<script src="codemirror/addon/searchcursor.js" type="text/javascript"></script>
<script src="codemirror/addon/show-hint.js" type="text/javascript"></script>
<script src="codemirror/addon/dialog.js" type="text/javascript"></script>
<script src="codemirror/lib/python.js" type="text/javascript"></script>
<script src="codemirror/addon/anyword-hint.js" type="text/javascript"></script>
<script src="static/bootstrap-3.0.0/js/bootstrap.min.js" type="text/javascript"></script>
<script src="static/bootstrap-3.0.0/js/bootstrap-switch.min.js" type="text/javascript"></script>
<script src="static/js/PythonHelper.js" type="text/javascript"></script>




<!--Styling -->
<link href='https://fonts.googleapis.com/css?family=Raleway' rel='stylesheet' type='text/css'>
<link href='https://fonts.googleapis.com/css?family=Inconsolata' rel='stylesheet' type='text/css'>
<link rel="stylesheet" href="codemirror/lib/codemirror.css">
<link rel="stylesheet" href="codemirror/theme/zenburn.css">
<link rel="stylesheet" href="codemirror/addon/dialog.css">
<link rel="stylesheet" href="codemirror/addon/matchesonscrollbar.css">
<link rel="stylesheet" href="codemirror/addon/show-hint.css">
<link rel="stylesheet" href="static/css/activecode.css">
<link rel="stylesheet" href="static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
<link rel="stylesheet" href="static/bootstrap-3.0.0/css/bootstrap-switch.min.css" type="text/css" />
<link rel="stylesheet" href="static/css/pytutor.css" type="text/css" />
<link rel="stylesheet" href="static/css/jquery-ui-1.10.3.custom.min.css" type="text/css" />
<link rel="stylesheet" href="static/css/custom_bootstrap.css" type="text/css" />

<!-- OpenGraph Protocol -->

<meta property='og:image' content='static/image/PythonHelper.png'/>

<meta name=viewport content="width=device-width, initial-scale=1">
<!-- Config For Rune-->
<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.useRunestoneServices = true;
</script>

<title>Python Helper</title>
</head> 
<body>
<a href="https://github.com/ethanchewy/PythonHelper"><img style="position: absolute; top: 0; right: 0; border: 0;" src="https://camo.githubusercontent.com/a6677b08c955af8400f44c6298f40e7d19cc5b2d/68747470733a2f2f73332e616d617a6f6e6177732e636f6d2f6769746875622f726962626f6e732f666f726b6d655f72696768745f677261795f3664366436642e706e67" alt="Fork me on GitHub" data-canonical-src="https://s3.amazonaws.com/github/ribbons/forkme_right_gray_6d6d6d.png"></a>

<div>
    <div class="text-center">
        <h1>Python Helper</h1>
        <p><i>In browser implementation of a python text editor with live syntax checking, visualizer, and compiler.</i></p>
        <p>Have Suggestions? Please fill out this 2 question survey: <a href="https://docs.google.com/forms/d/e/1FAIpQLSc9kVGglSHIVJ-TuBvzTAcR5fNnKQzmJaZxYoiPWWPJiHdguQ/viewform" target="_blank">Survey</a></p>
    </div>
    <!-- Text Editor -->
    <div class="row" id="core">
      <div class="col-md-8">
        <h3>Text Editor <button type="button" class="btn btn-info features-container" data-toggle="collapse" data-target="#demo">Features</button></h3>
        <div id="demo" class="collapse">
          <ul style="float:left;margin: 0 20px 0 0;">
            <li>Auto Line Highlight</li>
            <li>Search Functionality: Ctrl-F</li>
            <li>AutoComplete: Ctrl-[Space]</li>
          </ul>
          <ul style="float:left;margin: 0 20px 0 0;">
            <li>Grader</li>
            <li>History Navigation</li>
            <li>Error Log</li>
          </ul>
        </div>
        <div style="margin-top:2px;float:left; clear: left;">
          <p id="TryOut"><i>Click on the Editor to Try it Out!</i></p>
        </div>
        <textarea id="code" data-component="activecode" id=ac_example1 data-lang="python"    data-timelimit=25000  data-codelens="false" data-audio=''     data-caption='factorial'>
def factorial(n):
  return 1 if (n < 1) else n * factorial(n-1)
print factorial(4)

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(factorial(4),24,"Correct")

myTests().main()
        </textarea>
    </div>
    <!-- Grader -->
    <div class="col-md-4" id="grader_container">
      <h3>Grader <i>(In Progress) </i><input id="grader_check" type="checkbox" name="my-checkbox" size="normal" data-on-color="success" data-on-color="info" data-on-text="Show"  data-off-text="Close" checked></h3>
      <div id="grader-content">
        <div id="grader"></div>
        <h4>How does it work?</h4>
        
          <p>Example Grader:</p>
          <pre>
  from unittest.gui import TestCaseGui

  class myTests(TestCaseGui):
      def testOne(self):
          self.assertEqual(factorial(4),24,"Correct")

  myTests().main()
          </pre>
          <h4> Ideas for Next Iterations:</h4>
          <ul>
            <li>Create seperate text box for grader code</li>
            <li>Eliminate grader code and subsitute it with a form</li>
            <li>...</li>
          </ul>
      </div>
    </div>
    <!-- Errors -->
    <div class="col-md-8">
      <div id="errors"></div>
    </div>
    <!-- Current & Upcoming Features -->
    <div class="col-md-4" id="features-container">
      <h3>Features <input id="features" type="checkbox" name="my-checkbox" size="normal" data-on-color="success" data-on-color="info" data-on-text="Show"  data-off-text="Close" checked></h3>
      <div id="features-content">
      <h4>Upcoming Features:</h4>     
      <ul>
        <li>Live Syntax Checker</li>
        <li>Dynamic Editable Hint System</li>
        <li>Dynamic Python Helper Integration</li>
        <li>Make embedabble</li>
        <li>Time Limit (for educators)</li>
        <li>...</li>
      </ul>
      <h4>Have any ideas for features?</h4>
      Please fill out this 2 question survey: <a href="https://docs.google.com/forms/d/e/1FAIpQLSc9kVGglSHIVJ-TuBvzTAcR5fNnKQzmJaZxYoiPWWPJiHdguQ/viewform" target="_blank">Survey</a></p>
      </div>
    </div>
    
  </div>
</div>
<footer class="footer">
  <div class="container">
      Uses <a href="http://www.skulpt.org/">Skulpt</a>, <a href="https://github.com/RunestoneInteractive/RunestoneComponents">Runestone Components</a>, <a href="http://codemirror.net/">Codemirror</a>, and <a href="http://getbootstrap.com/">Bootstrap</a>
    </div>
</footer>
</body> 

</html> 